### Logical operators

```julia
false && true # AND
```

```julia
false || true # OR
```

Note that they "short-circuit!"

```julia
x = -42
x > 0 || error("x must be positive")
```

Output:
```
[7] invokelatest at .\essentials.jl:709 [inlined]
 [8] eventloop(::ZMQ.Socket) at C:\Users\Charles\.julia\packages\IJulia\a1SNk\src\eventloop.jl:8
 [9] (::IJulia.var"#15#18")() at .\task.jl:356
```

### Comparisons

```julia
1 == 1.0 # Equality
```

```julia
1 === 1.0 # Programmatically identical
```

```julia
3 < π
```

```julia
1 <= 1
```

```julia
.1 + .2
```

```julia
.1 + .2 ≈ .3 # \approx + TAB
```

Comparisons "chain"

Try inserting parentheses around one of these comparisons

```julia
2 == 2.0 == 0x02
```

```julia
x = 42
0 < x < 100 || error("x must be between 0 and 100")
```

# Assignment

Assignment in Julia is done with the single `=`. All it does is associates a name (on the left) to a value (on the right).

```julia
x = 1 # Use the name `x` for the value `1`
```

```julia
y = x # Use the name `y` for the value that `x` refers to
```

```julia
x = "hello!" # Decide you have a better use for the name `x`
```

```julia
y # Is still the value 1
```

```julia
ϵ = eps(1.0) # You can make your own unicode names
```

```julia
5ϵ # Juxtaposition is multiplication
```

## Updating operators

All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia
y += 1
```

```julia
# This is exactly the same as:
y = y + 1
```

Note that it's just re-purposing the _same name_ for a new value. This means that the type might even change:

```julia
y /= 2
```

# Strings

```julia
s1 = "I am a string."
```

```julia
s2 = """I am also a string. """
```

```julia
"Here, we get an "error" because it's ambiguous where this string ends "
```

```julia
"""Look, Mom, no "errors"!!! """
```

```julia
    println("""The other nice thing about triple-quoted
               string literals is that they ignore leading
               indentation, which is nice for long strings
               in real code. Try changing these quotes!""")
```

Strings are not written with single `'`s — that's used for a single character:

```julia
first(s1)
```

```julia
'⊂'
```

```julia
'If you try writing a string in single-quotes, you will get an error'
```

## String interpolation

You can use the dollar sign inside a string to evaluate a Julia expression inside a string — either a single variable or a more complicated expression:

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

```julia
 println("That is $(num_fingers + num_toes) digits in all!!")
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```