### Logical operators

```julia
false && true # AND
```

```julia
false || true # OR
```

Note that they "short-circuit!"

```julia
x = -42
x > 0 || error("x must be positive")
```

Output:
```
[7] invokelatest at .\essentials.jl:709 [inlined]
 [8] eventloop(::ZMQ.Socket) at C:\Users\Charles\.julia\packages\IJulia\a1SNk\src\eventloop.jl:8
 [9] (::IJulia.var"#15#18")() at .\task.jl:356
```

### Comparisons

```julia
1 == 1.0 # Equality
```

```julia
1 === 1.0 # Programmatically identical
```

```julia
3 < π
```

```julia
1 <= 1
```

```julia
.1 + .2
```

```julia
.1 + .2 ≈ .3 # \approx + TAB
```

Comparisons "chain"

Try inserting parentheses around one of these comparisons

```julia
2 == 2.0 == 0x02
```

```julia
x = 42
0 < x < 100 || error("x must be between 0 and 100")
```

## Updating operators

All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia
y += 1
```

```julia
# This is exactly the same as:
y = y + 1
```

Note that it's just re-purposing the _same name_ for a new value. This means that the type might even change:

```julia
y /= 2
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```